# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old.
    Ruth and Peter , their parents, have 3 kids."""

    # Matches uppercase letters followed by any number of lowercase letters which are then followed by
    # a space or a comma to find the name
    result = re.findall('[A-Z][a-z]+(?=[\s,])', simple_string)
    return result
    
names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
# assert len(names()) == 4, "There are four names in the simple_string"
```

```python
assert len(names()) == 4, "There are four names in the simple_string"
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        file = file.read()
        
    grades = re.findall('.+(?=:\sB)', file)
    return grades

grades()
```

Output:
```
Cell In[5], line 3, in grades()
      2 def grades():
----> 3     with open ("assets/grades.txt", "r") as file:
      4         file = file.read()
      6     grades = re.findall('.+(?=:\sB)', file)

File ~\AppData\Roaming\Python\Python311\site-packages\IPython\core\interactiveshell.py:324, in _modified_open(file, *args, **kwargs)
    317 if file in {0, 1, 2}:
    318     raise ValueError(
    319         f"IPython won't let you open fd={file} by default "
    320         "as it is likely to crash IPython. If you know what you are doing, "
    321         "you can use builtins' open."
    322     )
--> 324 return io_open(file, *args, **kwargs)

FileNotFoundError: [Errno 2] No such file or directory: 'assets/grades.txt'
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re


def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
        logdata = logdata.split('\n')  # turn the file into a list with each entry at each index

    logs = list()
    for log in logdata[:-1]:  # loop through every log in file
        log_dict = dict()
        log_dict["host"] = str(re.findall('\d+[.]\d+[.]\d+[.]\d+', log)).strip("[']")
        user_name = str(re.findall('\s.+(?=\s\[)', log)).strip("[']")
        if user_name == ' - -': user_name = '-'
        else:
            user_name = user_name.replace(' - ', '')
        log_dict['user_name'] = user_name


        log_dict['time'] = str(re.findall('\[.+\]', log)).strip("[']")

        request = str(re.findall('(?<=").+"', log)).strip("[']")
        log_dict['request'] = request.replace('"', '')

        logs.append(log_dict)

    return logs

print(logs())
print(len(logs()))
print(logs()[0])


```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```

```python

```

```python

```